```python
from IPython.core.interactiveshell import InteractiveShell
InteractiveShell.ast_node_interactivity = "all"

# Reload all modules imported with %aimport
%load_ext autoreload
%autoreload 1
```

```python
import RNN_helper
%aimport RNN_helper

helper = RNN_helper.rnn_helper()
```

```python
import tensorflow as tf
from tensorflow.keras.utils import to_categorical
from sklearn.model_selection import train_test_split
import numpy as np
import os
```

## Load the data

```python
data_dir = np.load('./data/dataset.npz')
X = data_dir['X']
y = data_dir['y']
y = to_categorical(y)
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.1, random_state=42)
```

## Load models and histories

```python
model_name_1 = "LSTM_sparsecat"
model1 = helper.loadModel(model_name_1)
score1 = model1.evaluate(X_test, y_test, verbose=0)
history1 = helper.loadHistory(model_name_1)

model_name_2 = "Only_head"
model2 = helper.loadModel(model_name_2)
score2 = model2.evaluate(X_test, y_test, verbose=0)
history2 = helper.loadHistory(model_name_2)
```

## Data dimension test

<!-- grade_id: data_dimension_test, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
maxlen = 40
maxfeatures = 1000
num_cases = y.shape[1]
n_samples, timestamps, features = X.shape

assert timestamps == maxlen
assert features == maxfeatures

### END HIDDEN TESTS
```

### 1. Model structure test

<!-- grade_id: lstm_test, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

n_units = 128

## Model1 LSTM sparse categorical cross entropy
lstm_layer = model1.get_layer(name='lstm_1')
dense_layer = model1.get_layer(name='dense_head')

# lstm layer tests
assert lstm_layer.input.shape[1:] == tf.TensorShape([maxlen, maxfeatures])
assert lstm_layer.output.shape[1:] == tf.TensorShape([n_units])

# Dense layer tests
assert dense_layer.output.shape[1:] == num_cases

### END HIDDEN TESTS
```

<!-- grade_id: simple_classification_test, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS

## Model1 LSTM sparse categorical cross entropy
dense_layer = model2.get_layer(name='dense_head')

# lstm layer tests
assert dense_layer.input.shape[1:] == tf.TensorShape([maxfeatures])
assert dense_layer.output.shape[1:] == tf.TensorShape(num_cases)

### END HIDDEN TESTS
```

### 2. Training result test

<!-- grade_id: lstm_train_test, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert history1['acc'][-1] > 0.8
assert history1['val_acc'][-1] > 0.7
### END HIDDEN TESTS
```

<!-- grade_id: simple_classification_train_test, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert history2['acc'][-1] > 0.7
assert history2['val_acc'][-1] > 0.7
### END HIDDEN TESTS
```

### 4. Evaluate the model

<!-- grade_id: lstm_evaluate, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score1[1] > 0.7
### END HIDDEN TESTS
```

<!-- grade_id: simple_classification_evaluate, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score2[1] > 0.7
### END HIDDEN TESTS
```

```python

```