```python
import numpy as np
import pandas as pd
```

```python
# Load submitted file 
submit = pd.read_csv('./my_submit.csv')
test_grad = pd.read_csv('../resource/asnlib/test_grad.csv')
```

<!-- grade_id: model_evaluate, points: 100, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
score = np.sum(submit.overall == test_grad.overall) / len(submit)
assert score >= 0.5
### END HIDDEN TESTS
```

```python

```