```python
import numpy as np
import pandas as pd
```

```python
# Load submitted file 
submit = pd.read_csv('./my_submit.csv')
test_grad = pd.read_csv('../resource/asnlib/test_grad.csv')
```

<!-- grade_id: model_evaluate0, points: 50, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
score = np.sum(submit.overall.values == test_grad.overall.values) / len(submit)
assert score >= 0.4
### END HIDDEN TESTS
```

<!-- grade_id: model_evaluate1, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score >= 0.45
### END HIDDEN TESTS
```

<!-- grade_id: model_evaluate2, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score >= 0.5
### END HIDDEN TESTS
```

<!-- grade_id: model_evaluate3, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score >= 0.6
### END HIDDEN TESTS
```