```python
import numpy as np
import pandas as pd
import os
```

```python
# Load submitted file 
submit = pd.read_csv('./my_submit.csv')

DATA_DIR = './Data'
file_with_targets = 'test_grad.csv'
if not os.path.exists(DATA_DIR):
    DATA_DIR = '../resource/asnlib/'
test_grad = pd.read_csv(os.path.join(DATA_DIR, file_with_targets))
```

<!-- grade_id: model_evaluate0, points: 70, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
score = np.sum(submit.overall.values == test_grad.overall.values) / len(submit)
assert score >= 0.5
### END HIDDEN TESTS
```

<!-- grade_id: model_evaluate2, points: 20, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score >= 0.55
### END HIDDEN TESTS
```

<!-- grade_id: model_evaluate3, points: 10, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert score >= 0.6
### END HIDDEN TESTS
```

```python

```